A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Graph Analysis.

In this problem, we will use the `networkx` library to perform basic graph analysis.

```python
%matplotlib inline
import networkx as nx
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns

from nose.tools import assert_equal, assert_is_instance, assert_true, assert_false
from numpy.testing import assert_array_equal, assert_array_almost_equal
```

## Zachary's Karate Club

We will use [Zachary's Karete Club](https://networkx.github.io/documentation/networkx-1.10/reference/generated/networkx.generators.social.karate_club_graph.html) data set. You can read more about it on [Wikipedia](https://en.wikipedia.org/wiki/Zachary%27s_karate_club) or refer to the original paper [Zachary (1997)](http://www.jstor.org.proxy2.library.illinois.edu/stable/3629752).

According to Wikipedia:

<blockquote>
A social network of a karate club was studied by Wayne W. Zachary for a period of three years from 1970 to 1972. The network captures 34 members of a karate club, documenting 78 pairwise links between members who interacted outside the club. During the study a conflict arose between the administrator "John A" and instructor "Mr. Hi" (pseudonyms), which led to the split of the club into two. Half of the members formed a new club around Mr. Hi, members from the other part found a new instructor or gave up karate. Basing on collected data Zachary assigned correctly all but one member of the club to the groups they actually joined after the split.
</blockquote>

We can use `matplotlib` and `networkx.draw()` to visualize the social network of Zachary's Karate Club. (You do not have to create this figure, although I suggest creating a similar graph as an optinal exercise.)

![](https://raw.githubusercontent.com/UI-DataScience/info490-sp16/master/Week10/assignments/images/karate_all.png)

## Compute Node degrees

- Your task is to write a function named `get_degree` that takes a `networkx.Graph` instance and computes the degree of each node.

- The return value should be a dictionary of `{node: degree}` for each node. For example, when you print out the result:
  ```python
  >>> g = nx.karate_club_graph()
  >>> degree = get_degree(g)
  >>> print('Node  Degree')
  >>> for node in degree:
  ...    print('{0:4d} {1:7d}'.format(node, degree[node]))
  ```
  you should get
  ```
  Node  Degree
   0      16
   1       9
   2      10
   3       6
   4       3
   5       4
   6       4
   7       4
   8       5
   9       2
  10       3
  11       1
  12       2
  13       5
  14       2
  15       2
  16       2
  17       2
  18       2
  19       3
  20       2
  21       2
  22       2
  23       5
  24       3
  25       3
  26       2
  27       4
  28       3
  29       4
  30       4
  31       6
  32      12
  33      17
```

<!-- grade_id: get_degree_answer, points: 0, grade: false, solution: true -->
```python
def get_degree(graph):

    # YOUR CODE HERE
```

```python
g = nx.karate_club_graph()
degree = get_degree(g)
print('Node  Degree')
for node in degree:
    print('{0:4d} {1:7d}'.format(node, degree[node]))
```

<!-- grade_id: get_degree_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(degree, dict)

keys = list(degree.keys())
assert_equal(len(keys), 34)
assert_equal(set(keys), set(range(34)))

values = [
    16, 9, 10, 6, 3, 4, 4, 4, 5, 2,
    3, 1, 2, 5, 2, 2, 2, 2, 2, 3,
    2, 2, 2, 5, 3, 3, 2, 4, 3, 4,
    4, 6, 12, 17
]
assert_equal(len(degree.values()), len(values))
assert_equal(set(degree.values()), set(values))
```

We see that two people (Nodes 0 and 33) have well above the average number of degrees. Node 0 corresponds to Mr. Hi, and node 33 corresponds to John A (or "officer"). Node 32 has many friends but not as many as those two. In the following figure, we visualize this by coloring each node according to whether it is a friend of node 0 or node 33. 

![](https://raw.githubusercontent.com/UI-DataScience/info490-sp16/master/Week10/assignments/images/karate_friends.png)

By coloring the friends of node 0 and node 33, you can see that the club already had two subgroups based on social relationships. Let's compare the previous visualization with the visualization of the actual breakup.

![](https://raw.githubusercontent.com/UI-DataScience/info490-sp16/master/Week10/assignments/images/karate_after.png)

(You do not have to create these figures, although I suggest creating a similar graph as an optinal exercise.)

These two figures (before and after the actual breakup) look very similar. Thus, we see the power of network analysis: we can gain insight to the underlying data by analyzing the properties of graphs.

-----