A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 2. Graph Concepts

In this problem, we use the [NetworkX](https://networkx.github.io/) library to create and visualize a graph.

```python
%matplotlib inline
import networkx as nx
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns

from nose.tools import assert_equal, assert_is_instance, assert_true
from numpy.testing import assert_array_equal, assert_array_almost_equal
```

## Create a directed graph.

- Make a directed graph with 20 integer nodes, and edges connecting every even integer node to the next odd integer node. Start from zero (Zero is an even number).

- `get_directed_graph()` should return a `networkx.DiGraph()` instance. For example,

```python
def get_directed_graph():
    dg = nx.DiGraph()
    # YOUR CODE HERE
    return dg
```

- The [Introduction to Network Analysis: Graph Concepts](https://github.com/UI-DataScience/accy571-fa16/blob/master/Week11/notebooks/intro2na-gc.ipynb) has examples of creating and drawing different types of graphs.

- Here's a sample graph:
![](digraph.png)

<!-- grade_id: get_directed_graph_answer, points: 0, grade: false, solution: true -->
```python
def get_directed_graph():
    """
    Creates a directed graph with 20 integer nodes.
    Every even interger node is connected to the next odd integer node.
    
    Paramters
    ---------
    None
    
    Returns
    -------
    A networkx.Digraph instance.
    """
    
    # YOUR CODE HERE
    
    return dg
```

```python
dg = get_directed_graph()
```

<!-- grade_id: get_directed_graph_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(dg, nx.DiGraph)

nodes = np.array(dg.nodes())
assert_equal(nodes.size, 20)
assert_equal(nodes.min(), 0)
assert_equal(nodes.max(), 19)
assert_true(np.all(np.diff(nodes) == 1))

edges = np.array(dg.edges())
assert_equal(edges.shape[0], 10)
assert_equal(edges.shape[1], 2)
assert_equal(edges[:, 0].min(), 0)
assert_equal(edges[:, 0].max(), 18)
assert_true(np.all(np.diff(edges[:, 0]) == 2))
assert_equal(edges[:, 1].min(), 1)
assert_equal(edges[:, 1].max(), 19)
assert_true(np.all(np.diff(edges[:, 1]) == 2))
```

## Draw the directed graph.

- Use `matplotlib` to draw the directed graph created with `get_directed_graph()`.

- Note that our function returns a Matplotlib axes instance. Thus, you should pass the `matplotlib.Axes` instance to the `ax` parameter in [NetworkX's drawing function](https://networkx.github.io/documentation/networkx-1.10/reference/drawing.html). When you create an Axes instance and then pass this to the `draw_circular()` function, `NetworkX` uses your custom Axes instance instead of creating a new one. For example, you can do

```python
def draw_directed_graph(dg):
    fig, ax = plt.subplots()
    # YOUR CODE HERE
    nx.draw_circular(..., ax=ax)
    # YOUR CODE HERE
    return ax
```

<!-- grade_id: draw_directed_graph_answer, points: 0, grade: false, solution: true -->
```python
def draw_directed_graph(dg):
    """
    Draws the directed graph 'dg'.
    
    Paramters
    ---------
    dg: A networkx.DiGraph instance.
    
    Returns
    -------
    A matplotlib.Axes instance.
    """
    
    # YOUR CODE HERE
    
    return ax
```

```python
ax1 = draw_directed_graph(dg)
plt.show()
```

<!-- grade_id: draw_directed_graph_test, points: 10, grade: true, solution: false -->
```python
pos = np.array(list(nx.circular_layout(dg).values()))
assert_is_instance(ax1, mpl.axes.Axes)
assert_array_almost_equal(ax1.collections[0].get_offsets(), pos)
for idx, seg in enumerate(ax1.collections[1].get_segments()):
    line = np.vstack((pos[idx * 2], pos[idx * 2 + 1]))
    assert_array_almost_equal(seg, line)
assert_equal(len(ax1.collections[2].get_segments()), 10)
```

```python

```