A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Semantic Analysis.

In this problem, we use [Wordnet](https://en.wikipedia.org/wiki/WordNet) to compute word similarites.

```python
from nltk.corpus import wordnet as wn

from nose.tools import (
    assert_is_instance,
    assert_almost_equal
    )
```

The [Introduction to NLP: Sematic Analysis notebook](https://github.com/UI-DataScience/accy571-fa16/blob/master/Week11/notebooks/intro2nlp-sa.ipynb) demonstrates how we can use Wordnet to compute paths between two words. Let's see how the results change when we use a different set of words.

## Compute the path similarity for cat, dog, boy, and girl.

- Write a function named `get_path_similarity_between_boy_and_girl()` that uses Wordnet to compute the path similarity between the words "boy" and "girl".

<!-- grade_id: boy_girl_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_boy_and_girl():
    """
    Computes the path similarity between 'boy' and 'girl'.
    
    Returns
    -------
    A float.
    """
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: boy_girl_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_boy_and_girl(), float)
assert_almost_equal(get_path_similarity_between_boy_and_girl(), 0.16666666666666666)
```

- Write a function named `get_path_similarity_between_boy_and_cat()` that uses Wordnet to compute the path similarity between the words "boy" and "cat".

<!-- grade_id: boy_cat_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_boy_and_cat():
    '''
    Computes the path similarity between "boy" and "cat".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: boy_cat_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_boy_and_cat(), float)
assert_almost_equal(get_path_similarity_between_boy_and_cat(), 0.08333333333333333)
```

- Write a function named `get_path_similarity_between_boy_and_dog()` that uses Wordnet to compute the path similarity between the words "boy" and "dog".

<!-- grade_id: boy_dog_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_boy_and_dog():
    '''
    Computes the path similarity between "boy" and "dog".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: boy_dog_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_boy_and_dog(), float)
assert_almost_equal(get_path_similarity_between_boy_and_dog(), 0.14285714285714285)
```

- Write a function named `get_path_similarity_between_girl_and_girl()` that uses Wordnet to compute the path similarity between the words "girl" and "girl".

<!-- grade_id: girl_girl_answer, points: 0, grade: false, solution: true -->
```python
def get_path_similarity_between_girl_and_girl():
    '''
    Computes the path similarity between "girl" and "girl".
    
    Returns
    -------
    A float.
    '''
    
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: girl_girl_test, points: 5, grade: true, solution: false -->
```python
assert_is_instance(get_path_similarity_between_girl_and_girl(), float)
assert_almost_equal(get_path_similarity_between_girl_and_girl(), 1.0)
```

In the following code cell, we print out the results of the four functions.

```python
fmt_str = '{1} to {2}: {0:4.3f}'

print('Path Similarity:')
print(40*'-')
print(fmt_str.format(get_path_similarity_between_boy_and_girl(), 'boy', 'girl'))
print(fmt_str.format(get_path_similarity_between_boy_and_cat(), 'boy', 'cat'))
print(fmt_str.format(get_path_similarity_between_boy_and_dog(), 'boy', 'dog'))
print(fmt_str.format(get_path_similarity_between_girl_and_girl(), 'girl', 'girl'))
```

Does it make sense?

```python

```